Differential Regulation of Hepatic Macrophage Fate by Chi3l1 in MASLD

Author response:

The following is the authors’ response to the original reviews

Reviewer #1 (Public review):

The manuscript by Shan et al seeks to define the role of the CHI3L1 protein in macrophages during the progression of MASH. The authors argue that the Chil1 gene is expressed highly in hepatic macrophages. Subsequently, they use Chil1 flx mice crossed to Clec4F-Cre or LysM-Cre to assess the role of this factor in the progression of MASH using a high-fat, high-cholesterol diet (HFHC). They found that loss of Chil1 in KCs (Clec4F Cre) leads to enhanced KC death and worsened hepatic steatosis. Using scRNA seq, they also provide evidence that loss of this factor promotes gene programs related to cell death. From a mechanistic perspective, they provide evidence that CHI3L serves as a glucose sink and thus loss of this molecule enhances macrophage glucose uptake and susceptibility to cell death. Using a bone marrow macrophage system and KCs they demonstrate that cell death induced by palmitic acid is attenuated by the addition of rCHI3L1. While the article is well written and potentially highlights a new mechanism of macrophage dysfunction in MASH, there are some concerns about the current data that limit my enthusiasm for the study in its current form. Please see my specific comments below.

(1) The authors' interpretation of the results from the KC (Clec4F) and MdM KO (LysM-Cre) experiments is flawed. For example, in Figure 2 the authors present data that knockout of Chil1 in KCs using Clec4f Cre produces worse liver steatosis and insulin resistance. However, in supplemental Figure 4, they perform the same experiment in LysM-Cre mice and find a somewhat different phenotype. The authors appear to be under the impression that LysM-Cre does not cause recombination in KCs and therefore interpret this data to mean that Chil1 is relevant in KCs and not MdMs. However, LysM-Cre DOES lead to efficient recombination in KCs and therefore Chil1 expression will be decreased in both KCs and MdM (along with PMNs) in this line.

Therefore, a phenotype observed with KC-KO should also be present in this model unless the authors argue that loss of Chil1 from the MdMs has the opposite phenotype of KCs and therefore attenuates the phenotype. The Cx3Cr1 CreER tamoxifen inducible system is currently the only macrophage Cre strategy that will avoid KC recombination. The authors need to rethink their results with the understanding that Chil1 is deleted from KCs in the LysM-Cre experiment. In addition, it appears that only one experiment was performed, with only 5 mice in each group for both the Clec4f and LysM-Cre data. This is generally not enough to make a firm conclusion for MASH diet experiments.

We thank the reviewer for raising this important point regarding our data interpretation. We have carefully examined the deletion efficiency of Chi3l1 in primary Kupffer cells (KCs) from Lyz2∆Chil1 (LysM-Cre) mice. Our results show roughly a 40% reduction in Chi3l1 expression at both the mRNA and protein levels (Revised Manuscript, Figure S7B and C). Given this modest decrease, Chi3l1 deletion in KCs of Lyz2∆Chil1 mice was incomplete, which likely accounts for the phenotypic differences observed between Clec4f∆Chil1 and Lyz2∆Chil1 mice in the MASLD model.

Furthermore, we have increased the sample size in both the Clec4f- and LysM-Cre experiments to 9–12 mice per group following the HFHC diet, thereby strengthening the statistical power and reliability of our findings (Revised Figures 2 and S8).

(2) The mouse weight gain is missing from Figure 2 and Supplementary Figure 4. This data is critical to interpret the changes in liver pathology, especially since they have worse insulin resistance.

We thank the reviewer for this valuable comment. We have now included the mouse body weight data in the revised manuscript (Figure 2A, B and Figures S8A, B). Compared with mice on a normal chow diet (NCD), all groups exhibited progressive weight gain during HFHC diet feeding. Notably, Clec4f∆Chil1 mice gained significantly more body weight than Chil1fl/fl controls, whereas Lyz2∆Chil1 mice showed a similar weight gain trajectory to Chil1fl/fl mice under the same conditions.

(3) Figure 4 suggests that KC death is increased with KO of Chil1. However, this data cannot be concluded from the plots shown. In Supplementary Figure 6 the authors provide a more appropriate gating scheme to quantify resident KCs that includes TIM4. The TIM4 data needs to be shown and quantified in Figure 4. As shown in Supplementary Figure 6, the F4/80 hi population is predominantly KCs at baseline; however, this is not true with MASH diets. Most of the recruited MoMFs also reside in the F4/80 hi gate where they can be identified by their lower expression of TIM4. The MoMF gate shown in this figure is incorrect. The CD11b hi population is predominantly PMNs, monocytes, and cDC,2 not MoMFs (PMID:33997821). In addition, the authors should stain the tissue for TIM4, which would also be expected to reveal a decrease in the number of resident KCs.

We thank the reviewer for raising this critical point regarding the gating strategy and interpretation of KC death. We have now refined our flow cytometry gating based on the reviewer’s suggestion. Specifically, we analyzed TIM4 expression and attempted to identify TIM4low MoMFs populations in our model. However, we did not detect a distinct TIM4low population, likely because our mice were fed the HFHC diet for only 16 weeks and had not yet developed liver fibrosis. We therefore reason that MoMFs have not fully acquired TIM4 expression at this stage.

To improve our analysis, we referred to published strategies (PMID: 41131393; PMID: 32562600) and gated KCs as CD45+CD11b+F4/80hi TIM4hi and MoMFs as CD45+Ly6G-CD11b+F4/80low TIM4low/-. Using this approach, we observed a gradual reduction of KCs and a corresponding increase in MoMFs in WT mice, with a significantly faster loss of KCs in Chil1-/- mice (Revised Figure 4C, D; Figure S10A).

Furthermore, immunofluorescence staining for TIM4 combined with TUNEL or cleaved caspase-3 confirmed an increased number of dying KCs in Chil1-/- mice compared to WT following HFHC diet feeding (Revised Figure 4E; Figure S10B).

(4) While the Clec4F Cre is specific to KCs, there is also less data about the impact of the Cre system on KC biology. Therefore, when looking at cell death, the authors need to include some mice that express Clec4F cre without the floxed allele to rule out any effects of the Cre itself. In addition, if the cell death phenotype is real, it should also be present in LysM Cre system for the reasons described above. Therefore, the authors should quantify the KC number and dying KCs in this mouse line as well.

We thank the reviewer for raising this important point. During our study, we indeed observed an increased number of KCs in Clec4f-Cre mice compared to WT controls, suggesting that the Clec4f-Cre system itself may modestly affect KC homeostasis. To address this, we compared KCs numbers between Clec4f∆Chil1 and Clec4f-Cre mice and found that Clec4f∆Chil1 mice displayed a significant reduction in KCs numbers following HFHC diet feeding. Moreover, co-staining for TIM4 and TUNEL revealed a marked increase in KCs death in Clec4f∆Chil1 mice relative to Clec4f-Cre mice, indicating that the observed phenotype is attributable to Chil1 deletion rather than Cre expression alone. These data have been reported in our related manuscript (He et al., bioRxiv, 2025.09.26.678483; doi: 10.1101/2025.09.26.678483).

In addition, we quantified KCs numbers and KCs death in the Lyz2-Cre line. TIM4/TUNEL co-staining showed comparable levels of KCs death between Chil1fl/fl and Lyz2∆Chil1 mice (Revised Figure S11B). Consistently, flow cytometry analyses revealed no significant differences in KCs numbers between these two groups before (0 weeks) or after (20 weeks) HFHC diet feeding (Revised Figures S11C, D). As discussed in our response to Comment 1, this may be due to the incomplete deletion of Chi3l1 in KCs (<50%) in the Lyz2-Cre line, which likely attenuates the phenotype.

(5) I am somewhat concerned about the conclusion that Chil1 is highly expressed in liver macrophages. Looking at our own data and those from the Liver Atlas it appears that this gene is primarily expressed in neutrophils. At a minimum, the authors should address the expression of Chil1 in macrophage populations from other publicly available datasets in mouse MASH to validate their findings (several options include - PMID: 33440159, 32888418, 32362324). If expression of Chil1 is not present in these other data sets, perhaps an environmental/microbiome difference may account for the distinct expression pattern observed. Either way, it is important to address this issue.

We thank the reviewer for this insightful comment and agree that analysis of scRNA-seq data, including our own and those reported in the Liver Atlas as well as in the referenced studies (PMID: 33440159, 32888418, 32362324), indicates that Chil1 is predominantly expressed in neutrophils.

However, our immunofluorescence staining under normal physiological conditions revealed that Chi3l1 protein is primarily localized in Kupffer cells (KCs), as demonstrated by strong co-staining with TIM4 (Revised Figure 1E). In MASLD mouse models induced by HFHC or MCD diets, we observed that both KCs and monocyte-derived macrophages (MoMFs) express Chi3l1, with particularly high levels in MoMFs.

We speculate that the apparent discrepancy between scRNA-seq datasets and our in situ findings may reflect differences in cellular proportions and detection sensitivity. Since hepatic macrophages (particularly KCs and MoMFs) constitute a larger proportion of total liver immune cells compared with neutrophils, their contribution to total Chi3l1 protein levels in tissue staining may appear dominant, despite lower transcript abundance per cell in sequencing datasets. We have included a discussion of this point in the revised manuscript to clarify this distinction (Revised manuscript, page 8,line 341-350 ).

Minor points:

(1) Were there any changes in liver fibrosis or liver fibrosis markers present in these experiments?

We assessed liver fibrosis using Sirius Red staining and α-SMA Western blot analysis.

We found no induction of liver fibrosis in our HFHC-induced MASLD model (Revised Figure S1A, B), but a clear elevation of fibrosis markers in the MCD-induced MASH model (Revised Figure S6A, B).

(2) In Supplementary Figure 3, the authors do a western blot for CHI3L1 in BMDMs. This should also be done for KCs isolated from these mice. Does this antibody work for immunofluorescence? Staining liver tissue would provide valuable information on the expression patterns.

We have included qPCR and western blot for Chi3l1 in isolated primary KCs from Lyz2∆Chil1 mice. The data show a slight, non-significant reduction in both mRNA and protein levels in KCs (Revised Figure S7B, C). The immunofluorescence staining on liver tissue showed that Chi3l1 is more likely expressed in the plasma membranes of TIM4+ F4/80+ KCs both under NCD and HFHC diet (Revised Figure 1E).

(3) What is the impact of MASH diet feeding on Chil1 expression in KCs or in the liver in general?

In both our MASLD and MASH models, diet feeding consistently upregulates Chi3l1 in KCs or in the liver in general (Revised Figure 1F, G, S6C,D).

(4) In Figure S1 the authors show tSNE plots of various monocyte and macrophage genes in the liver. Are these plots both diets together? How do things look when comparing these markers between the STD and HFHC diet? The population of recruited LAMs seems very small for 16 weeks of diet. Moreover, Chil1 should also be shown on these tSNE plots as well.

Yes, these plots are both diets together. When compared separately, the core marker expression is consistent between NCD and HFHC diets. However, the HFHC diet induces a relative increase in KC marker expression within the MoMF cluster, suggesting phenotypic adaptation (Author response image 1A, below). Moreover, Chil1 expression on the t-SNE plot was shown (Author response image 1B, below). However, compared to lineage-specific marker genes, Chi3l1 expression is rather low.

Gene expression levels of lineage-specific marker genes in monocytes/macrophages clusters between NCD and HFHC diets.

(A) UMAP plots show the scaled expression changes of lineage-specific markers in KCs/monocyte/macrophage clusters from mice under NCD and HFHC diets. Color represents the level of gene expression. (B) UMAP plots show the scaled expression changes of Chil1 in KCs/monocyte/macrophage clusters from mice under NCD and HFHC diets. Color represents the level of gene expression.

(5) In Figure 5, the authors demonstrate that CHI3L1 binds to glucose. However, given that all chitin molecules bind to carbohydrates, is this a new finding? The data showing that CHI3L is elevated in the serum after diet is interesting. What happens to serum levels of this molecule in KC KO or total macrophage KO mice? Do the authors think it primarily acts as a secreted molecule or in a cell-intrinsic manner?

We thank the reviewer for these insightful comments, which helped us clarify the novelty of our findings.

(1) Novelty of CHI3L1-Glucose Binding:

While chitin-binding domains are known to interact with carbohydrate polymers, our key discovery is that CHI3L1 (YKL-40)—a mammalian chitinase-like protein lacking enzymatic activity—specifically binds to glucose, a simple monosaccharide. This differs fundamentally from canonical binding to insoluble polysaccharides such as chitin and reveals a potential role for CHI3L1 in monosaccharide recognition, linking it to glucose metabolism and energy sensing. We clarified this point in the revised manuscript (page 9, line374-379).

(2) Serum CHI3L1 in Knockout Models:

Consistent with the reviewer’s suggestion, serum Chi3l1 levels are altered in our knockout models:

KC-specific KO (Clec4fΔChil1): Under normal chow, serum CHI3L1 is markedly reduced compared to controls and remains lower following HFHC feeding (Author response image 2A, below), indicating that Kupffer cells are the main source of circulating CHI3L1 under basal and disease conditions.

Macrophage KO (Lyz2ΔChil1): No significant changes were observed between Chil1fl/fl and Lyz2ΔChil1 mice under either diet (Author response image 2B, below), likely due to minimal monocyte-derived macrophage recruitment in this HFHC model (see Revised Figure 4C,D).

(3) Secreted vs. Cell-Intrinsic Role:

CHI3L1 predominantly localizes to the KC plasma membrane, consistent with a secreted role, and its serum reduction in KC-specific knockouts supports the physiological relevance of its secreted role. While cell-intrinsic effects have been reported elsewhere, our current data do not address this in KCs and warrant future investigation.

Chi3l1 expression in serum before and after HFHC in CKO mice.

(A) Western blot to detect Chi3l1 expression in serum of Chil1fl/fl and Clec4fΔChil1 mice before and after 16 weeks’ HFHC diet. n=3 mice/group. (B) Western blot to detect Chi3l1 expression in serum of Chil1fl/fl and Lyz2ΔChil1 before and after 16 weeks’ HFHC diet. n=3 mice/group.

Reviewer #2 (Public review):

The manuscript from Shan et al., sets out to investigate the role of Chi3l1 in different hepatic macrophage subsets (KCs and moMFs) in MASLD following their identification that KCs highly express this gene. To this end, they utilise Chi3l1KO, Clec4f-CrexChi3l1fl, and Lyz2-CrexChi3l1fl mice and WT controls fed a HFHC for different periods of time.

Major:

Firstly, the authors perform scRNA-seq, which led to the identification of Chi3l1 (encoded by Chil1) in macrophages. However, this is on a limited number of cells (especially in the HFHC context), and hence it would also be important to validate this finding in other publicly available MASLD/Fibrosis scRNA-seq datasets. Similarly, it would be important to examine if cells other than monocytes/macrophages also express this gene, given the use of the full KO in the manuscript. Along these lines, utilisation of publicly available human MASLD scRNA-seq datasets would also be important to understand where the increased expression observed in patients comes from and the overall relevance of macrophages in this finding.

We thank the reviewer for this valuable suggestion and acknowledge the limited number of cells analyzed under the HFHC condition in our original dataset. To strengthen our findings, we have now examined four additional publicly available scRNA-seq datasets— two from mouse models and two from human MASLD patients (Revised Figure S3, manuscript page 4, line 164-172). Across these datasets, the specific cell type showing the highest Chil1 expression varied somewhat between studies, likely reflecting model differences and disease stages. Nevertheless, Chil1 expression was consistently enriched in hepatic macrophage populations, including both Kupffer cells and infiltrating macrophages, in mouse and human livers. Notably, Chil1 expression was higher in infiltrating macrophages compared to resident Kupffer cells, supporting its upregulation during MASLD progression. These additional analyses confirm the robustness and crossspecies relevance of our finding that macrophages are the primary Chil1-expressing cell type in the liver.

Next, the authors use two different Cre lines (Clec4f-Cre and Lyz2-Cre) to target KCs and moMFs respectively. However, no evidence is provided to demonstrate that Chil1 is only deleted from the respective cells in the two CRE lines. Thus, KCs and moMFs should be sorted from both lines, and a qPCR performed to check the deletion of Chil1. This is especially important for the Lyz2-Cre, which has been routinely used in the literature to target KCs (as well as moMFs) and has (at least partial) penetrance in KCs (depending on the gene to be floxed). Also, while the Clec4f-Cre mice show an exacerbated MASLD phenotype, there is currently no baseline phenotype of these animals (or the Lyz2Cre) in steady state in relation to the same readouts provided in MASLD and the macrophage compartment. This is critical to understand if the phenotype is MASLD-specific or if loss of Chi3l1 already affects the macrophages under homeostatic conditions.

We thank the reviewer for raising this important point.

(1) Chil1 deletion efficiency in Clec4f-Cre and Lyz2-Cre lines:

We have assessed the efficiency of Chil1 deletion in both Lyz2∆Chil1 and Clec4f∆Chil1 mice by evaluating mRNA and protein levels of Chi3l1. For the Lyz2∆Chil1 mice, we measured Chi3l1 expression in bone marrow-derived macrophages (BMDMs) and primary Kupffer cells (KCs). Both qPCR (for mRNA) and Western blotting (for protein) reveal that Chi3l1 is almost undetectable in BMDMs from Lyz2∆Chil1 mice when compared to Chil1fl/fl controls. In contrast, we observe no significant reduction in Chi3l1 expression in KCs from these animals (Revised Figure S7B, C), suggesting Chil1 is deleted in BMDMs but not in KCs in Lyz2-Cre line.

For the Clec4f∆Chil1 mice, both mRNA and protein levels of Chi3l1 are barely detectable in BMDMs and primary KCs when compared to Chil1fl/fl controls (Revised Figure S4B, C). However, we did observe a faint Chi3l1 band in KCs of Clec4f∆Chil1 mice, which we suspect is due to contamination from LSECs during the KC isolation process, given that the TIM4 staining for KCs was approximately 90%. Overall, Chil1 is deleted in both KCs and BMDMs in Clec4f-Cre line.

Notably, since we observed a pronounced MASLD phenotype in Clec4f-Cre mice but not in Lyz2-Cre mice, these findings further underscore the critical role of Kupffer cells in the progression of MASLD.

(2) Whether the phenotype is MASLD-specific or whether loss of Chi3l1 already affects the macrophages under homeostatic conditions: We now included phenotypic data of Clec4fΔChil1 mice (KC-specific KO) and Lyz2∆Chil1 mice (MoMFs-specific KO) fed with NCD 16w (Revised Figure 2A-F, S8A-F). Shortly speaking, there is no baseline difference between Chil1fl/fl and Clec4fΔChil1 or Lyz2∆Chil1 mice in steady state in relation to the same readouts provided in MASLD.

Next, the authors suggest that loss of Chi3l1 promotes KC death. However, to examine this, they use Chi3l1 full KO mice instead of the Clec4f-Cre line. The reason for this is not clear, because in this regard, it is now not clear whether the effects are regulated by loss of Chi3l1 from KCs or from other hepatic cells (see point above). The authors mention that Chi3l1 is a secreted protein, so does this mean other cells are also secreting it, and are these needed for KC death? In that case, this would not explain the phenotype in the CLEC4F-Cre mice. Here, the authors do perform a basic immunophenotyping of the macrophage populations; however, the markers used are outdated, making it difficult to interpret the findings. Instead of F4/80 and CD11b, which do not allow a perfect discrimination of KCs and moMFs, especially in HFHC diet-fed mice, more robust and specific markers of KCs should be used, including CLEC4F, VSIG4, and TIM4.

We thank the reviewer for raising this important point. We performed experiments in Clec4f∆Chil1 (KC-specific KO) model. The phenotype in these mice closely mirrors that of the full KO: we observed a significant reduction in KC numbers and a concurrent increase in KC cell death following an HFHC diet in Clec4f∆Chil1 mice post HFHC diet compared to Clec4f-cre mice. We have reported these data in the following related manuscript (Figure 6 D-G). This confirms that the loss of CHI3L1 specifically from KCs is sufficient to drive this effect.

Hyperactivated Glycolysis Drives Spatially-Patterned Kupffer Cell Depletion in MASLD Jia He, Ran Li, Cheng Xie, Xiane Zhu, Keqin Wang, Zhao Shan bioRxiv 2025.09.26.678483; doi: https://doi.org/10.1101/2025.09.26.678483

While other hepatic cells (e.g., neutrophils and liver sinusoidal endothelial cells) also express Chi3l1, our data indicate that KC-secreted Chi3l1 plays a dominant and cellautonomous role in maintaining KCs viability. The potential contribution of other cellular sources to this phenotype remains an interesting direction for future study.

We apologize for the lack of clarity in our initial immunophenotyping. We have revised the flow cytometry data to clearly show that KCs are rigorously defined as TIM4+ cells (Revised Figure 4C, D).

Additionally, while the authors report a reduction of KCs in terms of absolute numbers, there are no differences in proportions. Thus, coupled with a decrease also in moMF numbers at 16 weeks (when one would expect an increase if KCs are decreased, based on previous literature) suggests that the differences in KC numbers may be due to differences in total cell counts obtained from the obese livers compared with controls. To rule this out, total cell counts and total live CD45+ cell counts should be provided. Here, the authors also provide tunnel staining in situ to demonstrate increased KC death, but as it is typically notoriously difficult to visualise dying KCs in MASLD models, here it would be important to provide more images. Similarly, there appear to be many more Tunel+ cells in the KO that are not KCs; thus, it would be important to examine this in the CLEC4F-Cre line to ascertain direct versus indirect effects on cell survival.

We thank the reviewer for raising this important point. We have now included the total cell counts and total live CD45+ cell counts, which showed similar numbers between WT and Chil1-/- mice post HFHC diet (Figure 3A, below).

Moreover, we included cleavaged caspase 3 and TIM4 co-staining in WT and Chil1-/- mice before and after HFHC diets, which confirmed increased KCs death in Chil1-/- mice (Revised Figure S10B). We have compared KCs number and KCs death between Clec4fcre and Clec4f∆Chil1 mice under NCD and HFHC diet in the following manuscript (Figure 6 D-G). The data showed similar KCs number under NCD and reduced KCs number in Clec4f∆Chil1 mice compared to Clec4f-cre mice, which confirms direct effects of Chi3l1 on cell survival but not because of cre insertion.

Hyperactivated Glycolysis Drives Spatially-Patterned Kupffer Cell Depletion in MASLD Jia He, Ran Li, Cheng Xie, Xiane Zhu, Keqin Wang, Zhao Shan bioRxiv 2025.09.26.678483; doi: https://doi.org/10.1101/2025.09.26.678483

Number of total cells and total live CD45+ cells in liver of WT and Chil1-/- mice.

(A) Number of total cells and total live CD45+ cells/liver were statistically analyzed. n= 3-4 mice per group.

Finally, the authors suggest that Chi3l1 exerts its effects through binding glucose and preventing its uptake. They use ex vivo/in vitro models to assess this with rChi3l1; however, here I miss the key in vivo experiment using the CLEC4F-Cre mice to prove that this in KCs is sufficient for the phenotype. This is critical to confirm the take-home message of the manuscript.

We agree that it is essential to confirm the in vivo relevance of Chi3l1-mediated glucose regulation in Kupffer cells (KCs). Our data suggest that KCs undergo cell death not because they express Chi3l1 per se, but because they exhibit a glucose-hungry metabolic phenotype that makes them uniquely dependent on Chi3l1-mediated regulation of glucose uptake. To directly assess this mechanism in vivo, we injected 2-NBDG, a fluorescent glucose analog, into overnight-fasted and refed mice and quantified its uptake in hepatic KCs. Notably, Chi3l1-deficient KCs exhibited significantly increased 2-NBDG uptake compared with controls, and this effect was markedly suppressed by co-treatment with recombinant Chi3l1 (rChi3l1) (Revised Figure 6G, H). These findings demonstrate that Chi3l1 regulates glucose uptake by KCs in vivo, supporting our proposed mechanism that Chi3l1 controls KC metabolic homeostasis through modulation of glucose availability.

Minor points:

(1) Some key references of macrophage heterogeneity in MASLD are not cited: PMID: 32362324 and PMID: 32888418.

We thank the reviewer for highlighting these critical references and have included them in the introduction (Revised manuscript, page 2, line 64-73).

(2) In the discussion, Figure 3H is referenced (Serum data), but there is no Figure 3H. If the authors have this data (increased Chi3l1 in serum of mice fed HFHC diet), what happens in CLEC4F-Cre mice fed the diet? Is this lost completely? This comes back to the point regarding the specificity of expression.

We apologize for the mistake. It should be Figure 5F now in the revised version, in which serum Chi3l1 was significantly upregulated after HFHC diet. Moreover, under a normal chow diet (NCD), serum CHI3L1 is significantly lower in Clec4fΔChil1 mice compared to controls (Chil1fl/fl). Following an HFHC diet, levels increase in both genotypes but remain relatively lower in the KC-KO mice (please see Figure 2A above). This data strongly suggests that Kupffer Cells (KCs) are the primary source of serum CHI3L1 under basal conditions and a major contributor during MASLD progression.

Reviewer #3 (Public review):

This paper investigates the role of Chi3l1 in regulating the fate of liver macrophages in the context of metabolic dysfunction leading to the development of MASLD. I do see value in this work, but some issues exist that should be addressed as well as possible.

(1) Chi3l1 has been linked to macrophage functions in MASLD/MASH, acute liver injury, and fibrosis models before (e.g., PMID: 37166517), which limits the novelty of the current work. It has even been linked to macrophage cell death/survival (PMID: 31250532) in the context of fibrosis, which is a main observation from the current study.

We thank the reviewer for this insightful comment regarding the novelty of our findings. We agree that Chi3l1 has previously been linked to macrophage survival and function in models of liver injury and fibrosis (e.g., PMID: 37166517, 31250532). However, our study focuses specifically on the early stage of MASLD, prior to the onset of fibrosis, revealing a distinct mechanistic role for CHI3L1 in this context.

We demonstrate that CHI3L1 directly interacts with extracellular glucose to regulate its cellular uptake—a previously unrecognized biochemical function. Furthermore, we show that CHI3L1’s protective role is metabolically dependent, safeguarding glucose-dependent Kupffer cells (KCs) but not monocyte-derived macrophages (MoMFs). This metabolic dichotomy and the direct link between CHI3L1 and glucose sensing represent conceptual advances beyond previous studies of CHI3L1 in fibrotic or injury models.

(2) The LysCre-experiments differ from experiments conducted by Ariel Feldstein's team (PMID: 37166517). What is the explanation for this difference? - The LysCre system is neither specific to macrophages (it also depletes in neutrophils, etc), nor is this system necessarily efficient in all myeloid cells (e.g., Kupffer cells vs other macrophages). The authors need to show the efficacy and specificity of the conditional KO regarding Chi3l1 in the different myeloid populations in the liver and the circulation.

We thank the reviewer for this important comment and the opportunity to clarify both the efficiency and specificity of our conditional knockouts, as well as the differences from the study by Feldstein’s group (PMID: 37166517).

(1) Chil1 deletion efficiency in Clec4f-Cre and Lyz2-Cre lines:

We have assessed the efficiency of Chil1 deletion in both Lyz2∆Chil1 and Clec4f∆Chil1 mice by evaluating mRNA and protein levels of Chi3l1. For the Lyz2∆Chil1 mice, we measured Chi3l1 expression in bone marrow-derived macrophages (BMDMs) and primary Kupffer cells (KCs). Both qPCR (for mRNA) and Western blotting (for protein) reveal that Chi3l1 is almost undetectable in BMDMs from Lyz2∆Chil1 mice when compared to Chil1fl/fl controls. In contrast, we observe no significant reduction in Chi3l1 expression in KCs from these animals (Revised Figure S7B, C), suggesting that Chil1 is deleted in BMDMs but not in KCs in Lyz2-Cre line.

For the Clec4f∆Chil1 mice, both mRNA and protein levels of Chi3l1 are barely detectable in BMDMs and primary KCs when compared to Chil1fl/fl controls (Revised Figure S4B, C). However, we did observe a faint Chi3l1 band in KCs of Clec4f∆Chil1 mice, which we suspect is due to contamination from LSECs during the KC isolation process, given that the TIM4 staining for KCs was approximately 90%. Overall, Chil1 is deleted in both KCs and BMDMs in Clec4f-Cre line.

Notably, since we observed a pronounced MASLD phenotype in Clec4f-Cre mice but not in Lyz2-Cre mice, these findings further underscore the critical role of Kupffer cells in the progression of MASLD.

(2) Explanation for Differences from Feldstein et al. (PMID: 37166517):

Our findings differ from those reported by Feldstein’s group primarily due to differences in disease stage and model. We used a high-fat, high-cholesterol (HFHC) diet to model earlystage MASLD characterized by steatosis and inflammation without fibrosis (Revised Figure S1A,B). In this context, we observed KC death but minimal MoMF infiltration (Revised Figure 4D). Accordingly, deletion of Chi3l1 in MoMFs (Lyz2∆Chil1) had no measurable effect on insulin resistance or steatosis, consistent with limited MoMF involvement at this stage. In contrast, the Feldstein study employed a CDAA-HFAT diet that models later-stage MASH with fibrosis. In that setting, Lyz2∆Chil1 mice showed reduced recruitment of neutrophils and MoMFs, which likely underlies the attenuation of fibrosis and disease severity reported. Together, these data support a model in which KCs and MoMFs play temporally distinct roles during MASLD progression: KCs primarily drive early lipid accumulation and metabolic dysfunction, whereas MoMFs contribute more substantially to inflammation and fibrosis at later stages.

(3) The conclusions are exclusively based on one MASLD model. I recommend confirming the key findings in a second, ideally a more fibrotic, MASH model.

We thank the reviewer for this valuable suggestion to validate our findings in an additional MASH model. We have now included data from a methionine- and choline-deficient (MCD) diet–induced MASH model, which exhibits pronounced hepatic lipid accumulation and fibrosis (Revised Figure S6A,B). Consistent with our HFHC results, Clec4f∆Chil1 mice displayed exacerbated MASH progression in this model, including increased lipid deposition, inflammation, and fibrosis (Revised Figure S6E-G).These findings confirm that CHI3L1 deficiency in Kupffer cells promotes hepatic lipid accumulation and disease progression across distinct MASLD/MASH models.

(4) Very few human data are being provided (e.g., no work with own human liver samples, work with primary human cells). Thus, the translational relevance of the observations remains unclear.

We thank the reviewer for this important comment regarding translational relevance. We fully agree that validation in human liver samples would further strengthen our study. However, obtaining tissue from early-stage steatotic livers is challenging due to the asymptomatic nature of this disease stage. Nonetheless, multiple studies have consistently reported Chi3l1 upregulation in human fibrotic and steatotic liver disease (PMID: 31250532, 40352927, 35360517), supporting the clinical significance of our mechanistic findings. We have now expanded the Discussion to highlight these human data and better contextualize our results within the spectrum of human MASLD/MASH progression (Revised manuscript, page 9, line390-394).

Minor points:

The authors need to follow the new nomenclature (e.g., MASLD instead of MAFLD, e.g., in Figure 1).

"MASLD" used throughout.

We thank the reviewers for their rigorous critique again. We thank eLife for fostering an environment of fairness and transparency that enables authors to communicate openly and present their data honestly.

Reference

(1) Tran, S. Baba I, Poupel L, et al(2020) Impaired Kupffer Cell Self-Renewal Alters the Liver Response to Lipid Overload during Non-alcoholic Steatohepatitis. Immunity 53, 627-640.

Changes from the earlier version to the revised version

Removed in the revised version: Abstract

Abstract

Metabolic dysfunction-associated steatotic liver disease (MASLD) progression involves the replacement of protective embryo-derived Kupffer cells (KCs) by inflammatory monocyte-derived macrophages (MoMFs), yet the regulatory mechanisms remain unclear. Here, we identify chitinase 3-like 1 (Chi3l1/YKL-40) as a critical metabolic regulator of hepatic macrophage fate. We observed high expression of Chi3l1 in both KCs and MoMFs during MASLD development. Genetic deletion of Chi3l1 specifically in KCs significantly exacerbated MASLD severity and metabolic dysfunction, whereas MoMF-specific Chi3l1 deletion showed minimal metabolic effects. Mechanistic studies revealed that this cell type-specific regulation arises from differential metabolic requirements: KCs display elevated glucose metabolism compared to MoMFs. Chi3l1 directly interacts with glucose to inhibit its cellular uptake, thereby selectively protecting glucose-dependent KCs from metabolic stress-induced cell death while having negligible effects on less glucose-dependent MoMFs. These findings uncover a novel Chi3l1-mediated metabolic checkpoint that preferentially maintains KCs populations through glucose metabolism modulation, providing important new insights into the pathogenesis of MASLD and potential therapeutic strategies targeting macrophage-specific metabolic pathways.

Introduction

Earlier version: Metabolic dysfunction-associated steatotic liver disease (MASLD) has become the most prevalent chronic liver disorder in western populations, affecting approximately 30% of adults and driven by its strong association with obesity and metabolic syndrome1. The disease spectrum ranges from matabolic dysfunction-associated fatty liver (MAFL) to metabolic dysfunction-associated steatohepatitis (MASH), with the latter characterized by steatosis, inflammation, hepatocyte ballooning, and progressive fibrosis2. Central to MASLD pathogenesis are hepatic macrophages, particularly the embryo-derived Kupffer cells (KCs) that reside in liver sinusoids3,4. These self-renewing resident macrophages5, play crucial roles in lipid homeostasis, as evidenced by studies showing that depletion of CD207+ KCs leads to impaired triglyceride storage6. When KCs die during MASH progression, they are gradually replaced by monocyte-derived macrophages (MoMFs) that exhibit more inflammatory properties and contribute to liver injury6,7.This transition from protective embryo-derived KCs (EmKCs) to inflammatory monocyte-derived KCs (MoKCs) represents a critical juncture in disease progression, yet the mechanisms regulating this shift remain poorly understood.

Revised version: Metabolic dysfunction-associated steatotic liver disease (MASLD) has become the most prevalent chronic liver disorder in western populations, affecting approximately 30% of adults and driven by its strong association with obesity and metabolic syndrome1. The disease spectrum ranges from matabolic dysfunction-associated fatty liver (MAFL) to metabolic dysfunction-associated steatohepatitis (MASH), with the latter characterized by steatosis, inflammation, hepatocyte ballooning, and progressive fibrosis2. Central to MASLD pathogenesis are hepatic macrophages, particularly the embryo-derived Kupffer cells (KCs) that reside in liver sinusoids3,4. These self-renewing resident macrophages5, play crucial roles in lipid homeostasis, as evidenced by studies showing that depletion of CD207+ KCs leads to impaired triglyceride storage6. As MASH progresses, dying KCs are progressively replaced by monocyte-derived macrophages (MoMFs) that exhibit heightened inflammatory properties and contribute to liver damage6,7. For example, one study demonstrated that in diet-induced MASH, KCs enhancer landscapes and gene expression profiles are profoundly reprogrammed (including up-regulation of Trem2 and Cd9) and KCs identity is lost, while MoMFs adopt convergent epigenomes, transcriptomes and functions during macrophage recruitment and adaptation in MASH8. Another work showed that in MASLD the number of resident KCs declines and MoMFs accumulate; these recruited macrophages include subsets that either mirror homeostatic KCs or resemble lipid-associated macrophages (LAMs) from obese adipose tissue, with the LAM-type expressing osteopontin and localizing to fibrotic zones9. Together, these findings highlight that this transition from protective embryo-derived KCs (EmKCs) to inflammatory monocyte-derived KCs (MoKCs) represents a critical juncture in disease progression, yet the mechanisms regulating this shift remain poorly understood.

Earlier version: A key determinant of macrophage function is cellular metabolism. Macrophages dynamically switch between glycolytic and oxidative phosphorylation pathways to adapt to environmental changes8. During MASLD, hepatic macrophages increase their glycolytic activity, which may exacerbate inflammation and tissue damage9-11. While glucose metabolism is known to influence macrophage polarization, its specific role in determining hepatic macrophage fate - particularly the balance between KCs and MoMFs - remains unknown. Chitinase 3-like 1 (Chi3l1/YKL-40) has emerged as an important regulator of macrophage biology, promoting cell survival through ERK1/2 and PI3K/Akt pathways while modulating anti-inflammatory cytokines like IL-1012-16. However, its potential role in macrophage metabolic reprogramming, particularly in the context of hepatic glucose metabolism, has not been explored.

Revised version: A key determinant of macrophage function is cellular metabolism. Macrophages dynamically switch between glycolytic and oxidative phosphorylation pathways to adapt to environmental changes10. During MASLD, hepatic macrophages increase their glycolytic activity, which may exacerbate inflammation and tissue damage11-13. While glucose metabolism is known to influence macrophage polarization, its specific role in determining hepatic macrophage fate - particularly the balance between KCs and MoMFs - remains unknown. Chitinase 3-like 1 (Chi3l1/YKL-40) has emerged as an important regulator of macrophage biology, promoting cell survival through ERK1/2 and PI3K/Akt pathways while modulating anti-inflammatory cytokines like IL-1014-18. However, its potential role in macrophage metabolic reprogramming, particularly in the context of hepatic glucose metabolism, has not been explored.

In this study, we identify a novel mechanism by which Chi3l1 governs hepatic macrophage fate through metabolic regulation. We demonstrate that Chi3l1 directly interacts with glucose to suppress its uptake in macrophages. Strikingly, this interaction selectively protects glucose-high KCs from cell death in MASLD conditions, while having minimal effect on glucose-low MoMFs. These findings reveal a previously unrecognized Chi3l1-mediated metabolic checkpoint that maintains KC populations, providing new insights into the pathogenesis of MASLD and potential therapeutic strategies.

Materials and methods

Animal experiments and procedures

Earlier version: Animals Chil1-/- (strain no. T014402), Chil1flox//flox (strain no. T013652), Lyz2-cre (strain no. T003822) and Clec4f-cre (strain no. T036801) mice with a C57BL/6J background were purchased from GemPharmatech. Rosa tdtomato mice (strain no. C001181) were purchased from Cyagen. Accordingly, C57BL/6J mice (strain no. N000013) were used as wild-type (WT) mice. To generate Clec4fΔChil1 mice, Chil1flox//flox mice were crossed with Clec4f-cre mice and knock out efficiency was examined in KCs by qRT-PCR (Figure S2B). To generate Clec4fRosa tdtomato mice, Rosa tdtomato mice were crossed with Clec4f-cre mice to examine the expression specificity of Clec4f-cre (Figure S2C). To generate Lyz2ΔChil1 mice, Chil1flox//flox mice were crossed with Lyz2-cre mice and knock out efficiency was examined in BMDM by western blot (Figure S3B). All mouse colonies were maintained at the Animal Core Facility of Yunnan University. The animal studies were approved by the Yunnan University Institutional Animal Care and Use Committee (IACUC, Approval No. YNU20220314). Male mice aged 6-8 weeks were used in this study.

Revised version: Animals Chil1-/- (strain no. T014402), Chil1flox//flox (strain no. T013652), Lyz2-cre (strain no. T003822), Clec4f-cre (strain no. T036801) with a C57BL/6J background were purchased from GemPharmatech. Rosa tdtomato mice (strain no. C001181) were purchased from Cyagen. Accordingly, C57BL/6J mice (strain no. N000013) were used as wild-type (WT) mice. To generate Clec4f△Chil1 mice, Chil1flox//flox mice were crossed with Clec4f-cre mice and knock out efficiency was examined in KCs or BMDM by qRT-PCR and western (Figure S4B, C). To generate Clec4fRosa tdtomato mice, Rosa tdtomato mice were crossed with Clec4f-cre mice to examine the expression specificity of Clec4f-cre (Figure S4D). To generate Lyz2△Chil1 mice, Chil1flox//flox mice were crossed with Lyz2-cre mice and knock out efficiency was examined in BMDM by western blot and qRT-PCR (Figure S7B, C). All mouse colonies were maintained at the Animal Core Facility of Yunnan University. The animal studies were approved by the Yunnan University Institutional Animal Care and Use Committee (IACUC, Approval No. YNU20220314). Male mice aged 6-8 weeks were used in this study.

Removed in the revised version: Construction of MASLD/MASH mouse model

Earlier version: Mice were provided a high-fat and high-cholesterol diet (Research Diet, d12108c, 40 kcal% fat and 1.25% cholesterol). Throughout the feeding period, the body weight and food consumption of the mice were observed and recorded weekly. Once the dietary intervention was completed, the mice were euthanized. Liver and murine serum samples were collected for further analysis. Alanine aminotransferase (ALT) and aspartate aminotransferase (AST) levels in the serum, as well as cholesterol (TC) and triglyceride (TG) levels in both serum and liver tissues, were quantified using commercially available kits (Nanjing Jiancheng Bioengineering Institute).

Revised version: Construction of MASLD/MASH mouse model Mice were provided a high-fat and high-cholesterol diet (Research Diet, d12108c, 40 kcal% fat and 1.25% cholesterol) or a methionine and choline deficient diet (Research Diet, A02082002BR) for 6 weeks. Throughout the feeding period, the body weight and food consumption of the mice were observed and recorded weekly. Once the dietary intervention was completed, the mice were euthanized. Liver and murine serum samples were collected for further analysis. Alanine aminotransferase (ALT) and aspartate aminotransferase (AST) levels in the serum, as well as cholesterol (TC) and triglyceride (TG) levels in both serum and liver tissues, were quantified using commercially available kits (Nanjing Jiancheng Bioengineering Institute).

Statistical analysis

Data are presented as mean ± standard error of the mean (SEM) in all graph figures. Statistical analyses were conducted using the SPSS statistics software (Version 22). To compare the two groups, an unpaired two-tailed Student’s t-test was used. One-way analysis of variance (ANOVA) was performed for comparisons involving three or more groups. For patients with MASLD liver, the samples were tested using the Mann-Whitney test. Statistical significance was set at p < 0.05 and p value is indicated. All cell culture results represent at least three independent experiments.

Additional Methods

Additional detailed methods can be found in the Supporting Information.

Results

Hepatic macrophages express Chi3l1

Earlier version: To investigate the dynamic changes in hepatic macrophages, we performed BD Rhapsody scRNA-seq on non-parenchymal cells (NPCs) isolated from healthy livers of mice fed a normal chow diet (NCD) and from livers of mice with MASLD induced by a high-fat, high-cholesterol (HFHC) diet. After quality control and filtration, we retained 23,312 cells from NCD livers and 6,567 cells from HFHC livers for downstream analysis. Using a graph-based clustering approach, we identified 32 distinct cell populations, visualized via uniform manifold approximation and projection (UMAP) (Figure 1A). Monocyte/macrophage subsets were further defined based on lineage-specific markers: Monocytes expressed Ly6c2, Chil3, S100a6, Ccr2, Itgam, and Cx3cr1 but lacked macrophage markers. KCs were marked by Cd68, Vsig4, Clec4f, Timd4, Adgre1, and Clec1b. MoMFs were negative for KCs markers but positive for macrophage markers such as Ccr2, Cx3cr1, Cd9, Itgax, Gpnmb, Cd68, and Adgre1 (Figure 1B, C; UMAP in Fig S1A)17,18. ScRNA-seq analysis revealed high expression of Chil1 (Chil1 being the gene name for Chi3l1) in hepatic macrophages (Figure 1D). To validate this, we performed immunofluorescence staining for Chi3l1 in hepatic macrophages using antibodies against Timd4, F4/80, and Chi3l1. Our results confirmed Chi3l1 expression in hepatic macrophages under both NCD and HFHC conditions (Figure 1E). Consistently, patients with MAFL or MASH exhibited elevated hepatic Chil1 mRNA levels, which correlated with MASLD severity and fibrosis stage (Figure 1F, G). These findings suggest that Chi3l1 is expressed in hepatic macrophages and may play a role in MASLD progression.

Revised version: To investigate the dynamic changes in hepatic macrophages, we performed BD Rhapsody scRNA-seq on non-parenchymal cells (NPCs) isolated from healthy livers of mice fed a normal chow diet (NCD) and from livers of mice with MASLD induced by a high-fat, high-cholesterol (HFHC) diet. Prior to NPC isolation for scRNA-seq, liver sections were subjected to H&E and Sirius Red staining, which revealed marked lipid accumulation without apparent fibrosis (Figure S1A). Consistently, western blot analysis of α-SMA showed no upregulation (Figure S1B), confirming that our HFHC model represents an early stage of MASLD. After quality control and filtration, we retained 23,312 cells from NCD livers and 6,567 cells from HFHC livers for downstream analysis. Using a graph-based clustering approach, we identified 32 distinct cell populations, visualized via uniform manifold approximation and projection (UMAP) (Figure 1A). Monocyte/macrophage subsets were further defined based on lineage-specific markers: Monocytes expressed Ly6c2, Chil3, S100a6, Ccr2, Itgam, and Cx3cr1 but lacked macrophage markers. KCs were marked by Cd68, Vsig4, Clec4f, TIM4, Adgre1, and Clec1b. MoMFs were negative for KCs markers but positive for macrophage markers such as Ccr2, Cx3cr1, Cd9, Itgax, Gpnmb, Cd68, and Adgre1 (Figure 1B, C; UMAP in Fig S2A)8,19. ScRNA-seq analysis revealed high expression of Chil1 (Chil1 being the gene name for Chi3l1) in hepatic macrophages (Figure 1D). To validate this, we performed immunofluorescence staining for Chi3l1 in hepatic macrophages using antibodies against TIM4, F4/80, and Chi3l1. Our results confirmed Chi3l1 expression in hepatic macrophages under both NCD and HFHC conditions (Figure 1E). To determine whether Chi3l1 expression is upregulated following HFHC diet feeding, we analyzed Chi3l1 protein levels in isolated KCs and whole liver tissue by western blotting. The results revealed a marked increase in Chi3l1 expression in both KCs and liver after HFHC diet (Figure 1F, G). To strengthen our findings, we further examined four additional publicly available scRNA-seq datasets—two from mouse models and two from human MASLD patients (Figure S3). Across these datasets, the specific cell type showing the highest Chil1 expression varied somewhat between studies, likely reflecting model differences and disease stages. Nevertheless, Chil1 expression was consistently enriched in hepatic macrophage populations, including both Kupffer cells and infiltrating macrophages, in mouse and human livers. Notably, Chil1 expression was higher in MoMFs compared to resident KCs, supporting its upregulation during MASLD progression. Consistently, patients with MAFL or MASH exhibited elevated hepatic Chil1 mRNA levels, which correlated with MASLD severity and fibrosis stage (Figure 1H, I). These findings suggest that Chi3l1 is expressed in hepatic macrophages and may play a role in MASLD progression.

Hepatic macrophages express Chi3l1.

Added in the revised version: (A-D) Wildtype C57BL/6J mice were fed either a normal chow diet (NCD) or HFHC for 16 weeks. NPCs were isolated and subjected to BD Rhapsody scRNA sequencing. (A) Uniform manifold approximation and projection (UMAP) plots illustrate the clustering of NPCs in the livers of mice fed NCD and HFHC. Cell clusters are color-coded, with monocytes/macrophages clusters outlined. (B) UMAP plots depict the clustering of Monocytes/Macrophages in the livers of mice fed NCD and HFHC. Cell clusters are color-coded. (C) Dot plot displays the scaled gene expression levels of lineage-specific marker genes in different cell clusters. (D) Dot plot shows the scaled gene expression levels of Chil1 in the indicated cell clusters. (E) Immunofluorescent staining of TIM4 (white), F4/80 (red), Chi3l1 (green), and nuclear DAPI (blue) in liver sections of mice fed with either NCD or HFHC for 16 weeks, illustrating Chi3l1 expression in hepatic macrophages. Scale bar=20μm and 5μm (zoom). (F, G) Western blot analysis of Chi3l1 in either isolated Kupffer cells (KCs, F) or whole liver tissue (Liver, G) from mice fed either NCD or HFHC diet. n=2-3 mice/group. (H) mRNA expression levels of Chil1 in liver tissues of patients with metabolic dysfunction-associated fatty liver (MAFL) or with metabolic dysfunction-associated steatohepatitis (MASH) (GEO Datasets: GSE167523, GSE207310, GSE130970). No-MAFLD or Healthy individuals serve as controls. (I) The correlation between mRNA expression levels of Chil1 and MASLD activity score or fibrosis stage was analyzed (GEO Datasets: GSE130970). Representative images were shown in E. Mann-Whitney test was performed in F. Pearson’s correlation was performed in G. P value and r value are as indicated.

Removed in the revised version: (A-D) Wildtype C57BL/6J mice were fed either a normal chow diet (NCD) or HFHC for 16 weeks. NPCs were isolated and subjected to BD Rhapsody scRNA sequencing.

Removed in the revised version: (A) Uniform manifold approximation and projection (UMAP) plots illustrate the clustering of NPCs in the livers of mice fed NCD and HFHC. Cell clusters are color-coded, with monocytes/macrophages clusters outlined.

Removed in the revised version: (B) UMAP plots depict the clustering of Monocytes/Macrophages in the livers of mice fed NCD and HFHC. Cell clusters are color-coded.

Removed in the revised version: (D) Dot plot displays the scaled gene expression levels of lineage-specific marker genes in different cell clusters.

Removed in the revised version: (E)Dot plot shows the scaled gene expression levels of Chil1 in the indicated cell clusters.

Removed in the revised version: (E) Immunofluorescent staining of Timd4 (white), F4/80 (red), Chi3l1 (green), and nuclear DAPI (blue) in liver sections of mice fed with either NCD or HFHC for 16 weeks, illustrating Chi3l1 expression in hepatic macrophages. Scale bar=20μm and 5μm (zoom).

Removed in the revised version: (G) mRNA expression levels of Chil1 in liver tissues of patients with metabolic dysfunction-associated fatty liver (MAFL) or with metabolic dysfunction-associated steatohepatitis (MASH) (GEO Datasets: GSE167523, GSE207310, GSE130970). No-MAFLD or Healthy individuals serve as controls.

Removed in the revised version: (H) The correlation between mRNA expression levels of Chil1 and MASLD activity score or fibrosis stage was analyzed (GEO Datasets: GSE130970).

Removed in the revised version: Representative images were shown in E. Mann-Whitney test was performed in F. Pearson’s correlation was performed in G. P value and r value are as indicated.

Deficiency of Chi3l1 in Kupffer cells promotes insulin resistance and hepatic lipid accumulation

Earlier version: Given that hepatic macrophages are the primary source of Chi3l1 in the liver, we investigated its functional role by generating mice with conditional knockout (cKO) of Chil1 in either KCs or MoMFs. First, we generated Clec4f ΔChil1 mice by crossing Chil1fl/fl mice with Clec4f-cre mice19, achieving selective Chil1 ablation in KCs (Figure S2A). When fed an HFHC diet, Clec4fΔChil1 mice exhibited: Increased hepatic lipid deposition, evident by H&E and Oil Red O staining as early as 8 weeks (Figure 2A, B). Elevated metabolic disturbances at 16 weeks, including higher liver index (liver/body weight), serum ALT, and increased cholesterol/triglycerides in both liver and serum (Figure 2C). Impaired glucose metabolism, with worsened glucose tolerance (IGTT) and insulin resistance (ITT) (Figure 2D, B). These results demonstrate that Chi3l1 deficiency in KCs drives lipid accumulation and insulin resistance.

Revised version: Given that Chi3l1 is highly expressed in hepatic macrophages, we investigated its functional role by generating mice with conditional knockout (cKO) of Chil1 in either KCs or MoMFs. First, we generated Clec4f △Chil1 mice by crossing Chil1fl/fl mice with Clec4f-cre mice20, achieving KC-specific deletion of Chil1 (Figure S4A-C). These mice, along with Chil1fl/fl controls, were fed either a NCD or a HFHC diet. Under NCD feeding, Clec4fΔChil1 and Chil1fl/fl mice displayed comparable phenotypes in terms of body weight gain, hepatic lipid deposition, metabolic parameters, glucose tolerance and insulin resistance (Figure 2A–F). In contrast, when fed an HFHC diet, Clec4fΔChil1 mice exhibited markedly accelerated weight gain compared to controls (Figure 2A,B). These mice also showed increased hepatic lipid accumulation, as evidenced by H&E and Oil Red O staining at 16 weeks (Figure 2C), along with greater metabolic disturbances, including a higher liver index (liver-to-body weight ratio), elevated serum ALT levels, and increased cholesterol and triglyceride levels in both liver and serum (Figure 2D). Furthermore, Clec4fΔChil1 mice exhibited impaired glucose metabolism, as indicated by worsened glucose tolerance and insulin resistance in IGTT and ITT assays (Figure 2E,F).To exclude potential off-target effects caused by Clec4f-Cre insertion, we compared Clec4f-Cre and Clec4fΔChil1 mice. The phenotypic differences between Clec4f-Cre and Clec4fΔChil1 mice mirrored those observed between Chil1fl/fl and Clec4fΔChil1 mice, with the latter showing faster weight gain, more severe hepatic steatosis, greater metabolic dysregulation, and worsened glucose intolerance and insulin resistance (Figure S5A–G).

Deficiency of Chi3l1 in Kupffer cells promotes insulin resistance and hepatic lipid accumulation.

Added in the revised version: Chil1fl/fl and Clec4f△Chil1 mice were fed either a normal chow diet (NCD) or a high-fat, high-cholesterol (HFHC) diet for 16 weeks. (A, B) Body weight was recorded during HFHC diet feeding (A) and expressed as a percentage of initial body mass (B). (C) H&E (Upper panel) and oil red o staining (Lower panel) was performed to examine liver histology and hepatic lipid accumulation in in both genotypes after 16 weeks of NCD or HFHC diet. Scale bar = 20 µm. (D) Liver index (liver weight/body weight), ALT levels, and serum and liver Cholesterol or Triglyceride levels were measured in both genotypes after 16 weeks on NCD or HFHC diets. n=4-12 mice/group. (E, F) Intraperitoneal glucose tolerance test (IGTT) and insulin tolerance test (ITT) were performed after 16 weeks of NCD or HFHC feeding in both genotypes (n = 4–12 mice per group). Representative images were shown in (A, E). One-way ANOVA was performed in (A, B, D-F). P-value is as indicated.

Added in the revised version: Additionally, we constructed another MASH model induced by a methionine and choline deficient diet (MCD), in which extensive inflammatory cell infiltration and fibrosis were induced(Figure S6A, B). Moreover, mRNA and protein levels of Chi3l1 were also upregulated in this model(Figure S6C,D). Although in this model body weight decrease instead of increase, Clec4f △Chil1 mice still showed severe hepatic steatosis and greater metabolic dysregulation (Figure S6E-G). These results further supported that Chi3l1 deficiency in KCs drives MASLD progression.

Added in the revised version: To assess the role of Chi3l1 in MoMFs, we generated Lyz2△Chil1 mice (Chil1fl/fl × Lyz2-Cre20) and validated Chil1 deletion efficiency in BMDM (Figure S7A–C). Chil1 expression was completely abolished in BMDM from Lyz2ΔChil1 mice. Considering the partial activity of Lyz2-Cre in KCs, we further assessed Chi3l1 expression in KCs isolated from Lyz2Chil1 mice. Only a modest (∼40%) reduction in Chi3l1 mRNA and protein levels was observed in KCs, indicating that Lyz2-Cre–mediated deletion minimally affects Chi3l1 expression in KCs. Lyz2ΔChil1 and Chil1fl/fl control mice were then fed either a NCD or a HFHC diet. Under both dietary conditions, the two genotypes exhibited comparable phenotypes with respect to body weight gain, hepatic lipid accumulation, metabolic parameters, glucose tolerance, and insulin sensitivity (Figure S8A–F). These results indicate that Chi3l1 loss in MoMFs does not substantially impact metabolic regulation.

Removed in the revised version: Chil1fl/fl and Clec4fΔChil1 mice were fed with a HFHC diet for 0, 8 and 16 weeks.

Removed in the revised version: (A) H&E staining was performed to examine liver histology of Chil1fl/fl and Clec4fΔChil1 mice fed the HFHC diet for 0, 8 and 16 weeks. Scale bar=20μm.

Removed in the revised version: (B) Oil red o staining was performed to detect fat accumulation in the livers of Chil1fl/fl and Clec4fΔChil1 mice fed the HFHC diet for 0, 8 and 16 weeks. Scale bar=20μm.

Removed in the revised version: (C) Liver index (liver weight/body weight), ALT levels, and serum and liver Cholesterol or Triglyceride levels were measured in Chil1fl/fl and Clec4fΔChil1 mice fed a HFHC diet for 16 weeks. n=5 mice/group.

Removed in the revised version: (D&E) i.p. glucose tolerance test (IGTT) and i.p. insulin tolerance test (ITT) were measured in Chil1fl/fl and Clec4fΔChil1 mice fed a HFHC diet for 16 weeks. n=5 mice/group.

Removed in the revised version: Representative images are shown in A, B. Two-tailed, unpaired student t-test was performed in C, D, E. P value is as indicated.

Removed in the revised version: To assess the role of Chi3l1 in MoMFs, we generated Lyz2ΔChil1 mice (Chil1fl/fl × Lyz2-Cre19) (Figure S3A). After 20 weeks on an HFHC diet: Hepatic lipid accumulation (H&E, Oil Red O) was comparable to controls (Figure S4A, B). Metabolic parameters (liver index, ALT, serum/liver lipids) showed no significant differences (Figure S4C). Lyz2ΔChil1 mice showed glucose intolerance (IGTT) but normal insulin sensitivity (ITT) (Figure S4D, E). Thus, Chi3l1 loss in MoMFs does not substantially contribute to metabolic dysregulation.

ScRNA-seq reveals upregulated glucose metabolism-related transcripts in KCs, correlating with cell death signatures

Earlier version: To dissect the distinct metabolic and functional profiles between KCs and MoMFs during MASLD progression, we analyzed our scRNA-seq data. Consistent with prior studies6,7, we observed decreased KCs numbers but increased MoMFs and monocytes in HFHC-fed mice compared to NCD controls (Figure 3A). Kyoto Encyclopedia of Genes and Genomes (KEGG) pathway analysis showed that while both cell types exhibited activation of phagocytosis-related pathways (lysosome, phagosome, endocytosis, and efferocytosis), they displayed divergent cell fate patterns (Figure 3B). KCs showed strong cell death signatures, whereas MoMFs maintained proliferative activity without evidence of cell death (Figure 3B). Monocytes showed strong cell death and proliferative activity (Figure 3B). Given the significant role of metabolic regulation in cell fate,20 we compared pathways involved in glucose metabolism, cell death, and cell proliferation (Figure 3C). Notably, glucose metabolism pathways were significantly more active in KCs and monocytes compared to MoMFs (Figure 3C). Moreover, the cell proliferation pathway was highly activated in monocytes and consistently activated in MoMFs but not in KCs (Figure 3C). Gene Set Variation Analysis (GSVA)-based correlation analysis revealed a striking association between glucose metabolism and cell death pathways (Figure 3D). These findings demonstrate distinct glucose metabolic activation patterns between KCs and MoMFs, which may underlie their divergent cell fates in MASLD progression.

Revised version: To dissect the distinct metabolic and functional profiles between KCs and MoMFs during MASLD progression, we analyzed our scRNA-seq data. Consistent with prior studies6,7, we observed decreased KCs numbers but increased MoMFs and monocytes in HFHC-fed mice compared to NCD controls (Figure 3A). Kyoto Encyclopedia of Genes and Genomes (KEGG) pathway analysis showed that while both cell types exhibited activation of phagocytosis-related pathways (lysosome, phagosome, endocytosis, and efferocytosis), they displayed divergent cell fate patterns (Figure 3B). KCs showed strong cell death signatures, whereas MoMFs maintained proliferative activity without evidence of cell death (Figure 3B). Monocytes showed strong cell death and proliferative activity (Figure 3B). Given the significant role of metabolic regulation in cell fate,21 we compared pathways involved in glucose metabolism, cell death, and cell proliferation (Figure 3C). Notably, glucose metabolism pathways were significantly more active in KCs and monocytes compared to MoMFs (Figure 3C). Moreover, the cell proliferation pathway was highly activated in monocytes and consistently activated in MoMFs but not in KCs (Figure 3C). Gene Set Variation Analysis (GSVA)-based correlation analysis revealed a striking association between glucose metabolism and cell death pathways (Figure 3D). These findings demonstrate distinct glucose metabolic activation patterns between KCs and MoMFs, which may underlie their divergent cell fates in MASLD progression.

ScRNA-seq reveals upregulated glucose metabolism-related transcripts in KCs, correlating with cell death signatures.

Removed in the revised version: (A-D) Wildtype C57BL/6J mice were fed either a normal chow diet (NCD) or HFHC for 16 weeks. NPCs were isolated and subjected to BD Rhapsody scRNA sequencing.

Removed in the revised version: (A) Quantification of each cell cluster is presented.

Removed in the revised version: (B) KEGG analysis reveals the top 12 enriched pathways for up-regulated genes when comparing HFHC versus NCD in KCs, monocytes, and MoMFs, respectively.

Earlier version: (C) Gene set variation analysis (GSVA) shows pathway activity for cell death, glucose metabolism, and cell proliferation in KCs, monocytes, and MoMFs of WT mice fed NCD or HFHC for 16 weeks, respectively.

Revised version: (A-D) Wildtype C57BL/6J mice were fed either a normal chow diet (NCD) or HFHC for 16 weeks. NPCs were isolated and subjected to BD Rhapsody scRNA sequencing. (A) Quantification of each cell cluster is presented. (B) KEGG analysis reveals the top 12 enriched pathways for up-regulated genes when comparing HFHC versus NCD in KCs, monocytes, and MoMFs, respectively. (C) Gene set variation analysis (GSVA) shows pathway activity for cell death, glucose metabolism, and cell proliferation in KCs, monocytes, and MoMFs of WT mice fed NCD or HFHC for 16 weeks, respectively. (D) The correlation between cell death and glucose metabolism pathways, based on GSVA score, is depicted.

Removed in the revised version: (D) The correlation between cell death and glucose metabolism pathways, based on GSVA score, is depicted.

Chi3l1 deficiency promote KCs death during MASLD

Earlier version: To investigate the role of Chi3l1 in KCs survival during MASLD, we generated Chil1-/- mice, as Chi3l1 is a secreted protein21,22. The successful knockout was confirmed by qRT-PCR analysis of liver tissue (Figure S5A, B). We then performed scRNA-seq on NPCs isolated from Chil1-/- mice fed an HFHC diet for 16 weeks. After quality control, 6,813 high-quality cells were retained for analysis. Using established KC markers (Figure 1C), we conducted GSVA to examine metabolic pathways. This revealed enhanced cell death pathways in KCs from HFHC-fed mice, with significantly greater apoptosis signatures in Chil1-/- KCs compared to wild-type (WT) controls (Figure 4A). The increased apoptosis was further supported by upregulation of pro-apoptotic genes in Chil1-/- KCs (Figure 4B).

Revised version: To investigate the role of Chi3l1 in KCs survival during MASLD, we generated Chil1-/- mice, as Chi3l1 is a secreted protein22,23. The successful knockout was confirmed by qRT-PCR analysis of liver tissue (Figure S9A, B). We then performed scRNA-seq on NPCs isolated from Chil1-/- mice fed an HFHC diet for 16 weeks. After quality control, 6,813 high-quality cells were retained for analysis. Using established KC markers (Figure 1C), we conducted GSVA to examine metabolic pathways. This revealed enhanced cell death pathways in KCs from HFHC-fed mice, with significantly greater apoptosis signatures in Chil1-/- KCs compared to wild-type (WT) controls (Figure 4A). The increased apoptosis was further supported by upregulation of pro-apoptotic genes in Chil1-/- KCs (Figure 4B).

Chi3l1 deficiency promote KCs death during MASLD.

Added in the revised version: (A) GSVA analysis showed the enrichment of cell death-related pathways in KCs from WT mice fed with either NCD or HFHC or Chil1-/- mice fed with HFHC. (B) Dot plot showing the scaled gene expression levels of Apoptosis-related genes and repressor genes in KCs from either WT or Chil1-/- fed with HFHC. (C) Strategy used to gate KCs (CD45+ F4/80hi CD11blow TIM4hi) and MoMFs (CD45+ F4/80low CD11bhi Ly6G− TIM4−) in the liver by flow cytometry. (D) Number of KCs and MoMFs /liver or gram(g) liver were statistically analyzed. n= 3-4 mice per group. (E) Immunofluorescent staining to detect TIM4(green), TUNEL (red), and nuclear DAPI (blue) in liver sections. Scale bar=20μm and 5μm(zoom). TUNEL+ TIM4+ cells/TIM4+ cells were statistically analyzed. n=4-6 mice/group. Representative images are shown in C, E. One-way ANOVA was performed in D. Two-tailed, unpaired student t-test was performed in E. P value is as indicated.

Removed in the revised version: (A) GSVA analysis showed the enrichment of cell death-related pathways in KCs from WT mice fed with either NCD or HFHC or Chil1-/- mice fed with HFHC.

Removed in the revised version: (B) Dot plot showing the scaled gene expression levels of Apoptosis-related genes and repressor genes in KCs from either WT or Chil1-/- fed with HFHC.

Removed in the revised version: (C) Flow cytometry analysis of CD45+ F4/80hi CD11blow Timd4+ Kupffer cells on nonparenchymal liver cells (NPC) between WT and Chil1-/- mice.

Removed in the revised version: (D) Number of KCs and MoMFs /liver or gram(g) liver were statistically analyzed. n= 3-4 mice per group.

Removed in the revised version: (E) Immunofluorescent staining to detect Clec4f (green), Tunel (red), and nuclear DAPI (blue) in liver sections. Scale bar=20μm and 5μm(zoom). Tunel+ cells/Clec4f+ cells were statistically analyzed. n=3-4 mice/group.

Removed in the revised version: Representative images are shown in C, E. One-way ANOVA was performed in D. Two-tailed, unpaired student t-test was performed in E. P value is as indicated.

Earlier version: Flow cytometry analysis validated these findings, using the gating strategy shown in Figure S6A. While WT and Chil1-/- mice showed similar KC numbers at baseline, dramatic differences emerged during HFHC feeding. WT KC numbers remained stable at 8 weeks but decreased by 50% at 16 weeks. In contrast, Chil1-/- mice exhibited accelerated KCs loss, with a 30% reduction by 8 weeks progressing to 60% by 16 weeks (Figure 4C, B). Notably, MoMFs populations remained comparable between groups at early timepoints but showed greater reduction in Chil1-/- mice at 16 weeks (Figure 4C, B).

Revised version: We next validated these findings by flow cytometry using the gating strategy shown in Figure 4C. While WT and Chil1-/- mice showed similar KC numbers at baseline, dramatic differences emerged during HFHC feeding. WT KC numbers remained stable at 8 weeks but decreased by 50% at 16 weeks. In contrast, Chil1-/- mice exhibited accelerated KCs loss, with a 30% reduction by 8 weeks progressing to 60% by 16 weeks (Figure 4D, S10A). Notably, MoMFs populations remained comparable between groups at early timepoints but showed greater reduction in Chil1-/- mice at 16 weeks ((Figure 4D, S10A).

Added in the revised version: Histological analysis further supported these findings. TIM4/TUNEL co-staining revealed no TUNEL+ KCs in WT livers at baseline, whereas 40% and 50% of KCs were TUNEL+ at 8 and 16 weeks, respectively. In Chil1-/- mice, KC apoptosis was significantly increased at both time points (Figure 4E). Consistent results were obtained with TIM4/cleaved caspase-3 co-staining (Figure S10B). We further confirmed these observations in Clec4fΔChil1 mice in both HFHC24 and MCD diet models. In the MCD model, Clec4fΔChil1 mice exhibited enhanced KCs death compared with Chil1fl/fl controls (Figure S11A). To exclude potential effects of myeloid cell–derived Chi3l1 on KCs survival, we compared KCs death and abundance between Chil1fl/fl and Lyz2ΔChil1 mice using histological and flow cytometric analyses. Loss of Chi3l1 in MoMFs did not lead to significant KC apoptosis or depletion (Figure S11B–D).Together, these results demonstrate that Chi3l1 deficiency promotes KC apoptosis, resulting in premature KC depletion during MASLD progression.

Removed in the revised version: Histological analysis through Clec4f/TUNEL co-staining provided further confirmation. While WT livers showed no TUNEL+ KCs at baseline, 40% were TUNEL+ by 8 weeks, increasing to 55% at 16 weeks. Chil1-/- mice displayed elevated KC apoptosis even at baseline, with significantly more TUNEL+ KCs (approximately 2-fold higher than WT) by 8 weeks (Figure 4E). Together, these results demonstrate that Chi3l1 deficiency accelerates HFHC-induced KCs death, leading to premature KCs depletion during MASLD progression.

Molecular interaction between Chi3l1 and glucose

Earlier version: Our investigation into Chi3l1-mediated KCs survival revealed an unexpected structural relationship: Chi3l1 binds to glucose, which is structurally analogous to chitin, a polysaccharide well known to bind Chi3l1(Figure 5A). Bioinformatics analysis using the STITCH database further supported this observation, predicting a high probability of direct Chi3l1-glucose interaction (Figure 5B). To experimentally validate this interaction, we performed pull-down assays using biotin-labeled glucose incubated with plasma from HFHC-fed mice. Streptavidin bead isolation followed by anti-Chi3l1 Western blotting demonstrated specific binding between Chi3l1 and biotin-glucose, but not biotin alone (Figure 5C, B). This interaction was competitively inhibited by unlabeled glucose, confirming specificity (Figure 5D). Quantitative analysis using microscale thermophoresis with recombinant mouse Chi3l1 (rChi3l1) yielded a dissociation constant (Kd) of 4.95 mM for the Chi3l1-glucose interaction (Figure 5E). Notably, circulating Chi3l1 levels were significantly elevated in serum from HFHC-fed mice compared to baseline (Figure 5F), suggesting a potential physiological role for this interaction in metabolic regulation. These findings establish Chi3l1 as a novel glucose-binding protein that may participate in glucose homeostasis during MASLD progression.

Revised version: Our investigation into Chi3l1-mediated KCs survival revealed an unexpected structural relationship: Chi3l1 binds to glucose, which is structurally analogous to chitin, a polysaccharide well known to bind Chi3l1(Figure 5A). Bioinformatics analysis using the STITCH database further supported this observation, predicting a high probability of direct Chi3l1-glucose interaction (Figure 5B). To experimentally validate this interaction, we performed pull-down assays using biotin-labeled glucose incubated with plasma from HFHC-fed mice. Streptavidin bead isolation followed by anti-Chi3l1 Western blotting demonstrated specific binding between Chi3l1 and biotin-glucose, but not biotin alone (Figure 5C, D). This interaction was competitively inhibited by unlabeled glucose, confirming specificity (Figure 5D). Quantitative analysis using microscale thermophoresis with recombinant mouse Chi3l1 (rChi3l1) yielded a dissociation constant (Kd) of 4.95 mM for the Chi3l1-glucose interaction (Figure 5E). Notably, circulating Chi3l1 levels were significantly elevated in serum from HFHC-fed mice compared to baseline (Figure 5F), suggesting a potential physiological role for this interaction in metabolic regulation. These findings establish Chi3l1 as a novel glucose-binding protein that may participate in glucose homeostasis during MASLD progression.

Molecular interaction between Chi3l1 and glucose.

Removed in the revised version: (A) A comparison of chemical structures between glucose and chitin.

Removed in the revised version: (B) Prediction of Chi3l1-glucose interaction using STITCH database (http://stitch.embl.de).

Removed in the revised version: (C) Strategy for pulling down glucose-binding proteins in murine serum.

Earlier version: (D) Biotin-conjugated glucose was incubated with murine serum from mice fed with HFHC for 16 weeks. Proteins bound to glucose were precipitated by streptavidin beads. Biotin or biotin-conjugated glucose plus glucose were used as negative controls. Western blot was performed to examine Chi3l1 in the precipitate.

Revised version: (A) A comparison of chemical structures between glucose and chitin. (B) Prediction of Chi3l1-glucose interaction using STITCH database (http://stitch.embl.de). (C) Strategy for pulling down glucose-binding proteins in murine serum. (D) Biotin-conjugated glucose was incubated with murine serum from mice fed with HFHC for 16 weeks. Proteins bound to glucose were precipitated by streptavidin beads. Biotin or biotin-conjugated glucose plus glucose were used as negative controls. Western blot was performed to examine Chi3l1 in the precipitate. (E) Microscale thermophoresis assay to detect the interaction between recombinant mouse Chi3l1 (rChi3l1) and glucose. Kd=4.95±0.66mM. (F) Western blot to detect Chi3l1 expression in murine serum before and after HFHC feeding. n=3 mice/group.

Removed in the revised version: (E) Microscale thermophoresis assay to detect the interaction between recombinant mouse Chi3l1 (rChi3l1) and glucose. Kd=4.95±0.66mM.

Removed in the revised version: (F) Western blot to detect Chi3l1 expression in murine serum before and after HFHC feeding. n=3 mice/group.

Chi3l1 limits glucose uptake and protects hepatic macrophages from cell death

Earlier version: To elucidate the functional consequences of Chi3l1-glucose binding, we examined glucose metabolism in hepatic macrophages. Using the fluorescent glucose analog 2-NBDG23, we performed uptake assays in KCs following 12-hour glucose starvation. While glycogen droplet size remained unchanged in untreated KCs regardless of rChi3l1 supplementation (Figure 6A), 2-NBDG exposure significantly increased glycogen accumulation. This effect was markedly suppressed by rChi3l1 co-treatment (Figure 6A), a phenotype replicated in BMDM (Figure 6B). These results demonstrate that Chi3l1 restricts glucose uptake and subsequent glycogen storage.

Revised version: To elucidate the functional consequences of Chi3l1-glucose binding, we examined glucose metabolism in hepatic macrophages. Using the fluorescent glucose analog 2-NBDG25, we performed uptake assays in KCs following 12-hour glucose starvation. While glycogen droplet size remained unchanged in untreated KCs regardless of rChi3l1 supplementation (Figure 6A), 2-NBDG exposure significantly increased glycogen accumulation. This effect was markedly suppressed by rChi3l1 co-treatment (Figure 6A), a phenotype replicated in BMDM (Figure 6A). These results demonstrate that Chi3l1 restricts glucose uptake and subsequent glycogen storage.

Chi3l1 limits glucose uptake and protects hepatic macrophages from cell death.

Added in the revised version: (A) Following 12 h of glucose starvation, isolated KCs or BMDM were divided into two groups: one treated with no 2-NBDG and the other with 2-NBDG. Within each group, KCs or BMDM were further treated without or with recombinant murine Chi3l1 (rChi3l1) for 6 h. Glycogen aggregate formation labeled by 2-NBDG (Green) in KCs or BMDM was examined after counterstaining with nuclear DAPI (Blue). Scale bar=2μm. Area of 2-NBDG in KCs were quantified. (B) Following 12 h of glucose starvation, BMDM were treated with either no glucose or high glucose (25mM). Concurrently, BMDM were treated without or with rChi3l1 for 24 h under each condition. glycogen aggregate formation in BMDM was detected using immunofluorescence staining for Stbd1 (red) and nuclear DAPI (blue). Scale bar = 10 μm. (C and D) BMDM cells were treated without or with rChi3l1 for 24 h and subjected to Seahorse metabolic analysis to measure the extracellular acidification rate (ECAR). (E and F) KCs were treated without (blank) or with either Isopropyl alcohol(Iso) or 800uM palmitic acid (PA) or 100ng rChi3l1 with 800 uM PA for 24 h. Western blot was performed to detect cleaved caspase 3 (Cl-Casp3) in E. Calcein/PI staining was quantified to detect cell viability in F. Scale bar=50μm. (G) Measurement of 2-NBDG (a fluorescent glucose analog) uptake by KCs in vivo. WT and Chil1-/- mice, either untreated or supplemented with rChi3l1, were injected intraperitoneally with 12 mg/kg 2-NBDG. After 45mins, KCs were isolated and glucose uptake assessed by spectrophotometry. (H) Representative immunofluorescence images of liver sections stained for TIM4 (red) and 2-NBDG uptake (green) to visualize glucose uptake by KCs in situ. Scale bar = 10 µm (zoom). Quantification is shown as the percentage of TIM4+ cells that are also 2-NBDG+. Representative images were shown in A, B, H. One-way ANOVA was performed in A, F, G, H. Two-tailed, unpaired student t-test was performed in D. P value is as indicated.

Removed in the revised version: (A) Following 12 h of glucose starvation, isolated KCs were divided into two groups: one treated with no 2-NBDG and the other with 2-NBDG. Within each group, KCs were further treated without or with recombinant murine Chi3l1 (rChi3l1) for 6 h. Glycogen aggregate formation labeled by 2-NBDG (Green) in BMDM was examined after counterstaining with nuclear DAPI (Blue). Scale bar=2μm.

Removed in the revised version: (B) Following 12 h of glucose starvation, BMDM were divided into two groups: one treated with no 2-NBDG and the other with 2-NBDG. Within each group, BMDM were further treated without or with recombinant murine Chi3l1 (rChi3l1) for 6 h. Glycogen aggregate formation labeled by 2-NBDG (Green) in BMDM was examined after counterstaining with nuclear DAPI (Blue). Scale bar=2μm. The area of 2-NBDG-labeled aggregate foci was quantified.

Removed in the revised version: (C) Following 12 h of glucose starvation, BMDM were treated with either no glucose or high glucose (25mM). Concurrently, BMDM were treated without or with rChi3l1 for 24 h under each condition. glycogen aggregate formation in BMDM was detected using immunofluorescence staining for Stbd1 (red) and nuclear DAPI (blue). Scale bar = 10 μm.

Removed in the revised version: (D and E) BMDM cells were treated without or with rChi3l1 for 24 h and subjected to Seahorse metabolic analysis to measure the extracellular acidification rate (ECAR).

Removed in the revised version: (F and G) KCs were treated without (blank) or with either Isopropyl alcohol(Iso) or 800uM palmitic acid (PA) or 100ng rChi3l1 with 800 uM PA for 24 h. Western blot was performed to detect cleaved caspase 3 (Cl-Casp3) in F. Calcein/PI staining was quantified to detect cell viability in G. Scale bar=50μm.

Removed in the revised version: Representative images are shown in A-C. One-way ANOVA was performed in B, G. Two-tailed, unpaired student t-test was performed in E. P value is as indicated.

Earlier version: Further validation using Stbd1 (a glycogen-binding protein23) immunofluorescence revealed minimal glycogen foci in glucose-deprived BMDM, with no rChi3l1-dependent differences. High-glucose conditions, however, triggered robust glycogen aggregation, which was significantly attenuated by rChi3l1 (Figure 6C). Concordantly, extracellular acidification rate (ECAR) measurements showed reduced basal and total glycolytic capacity in rChi3l1-treated BMDMs (Figure 6D,E), confirming Chi3l1’s role in limiting glucose metabolism.

Revised version: Further validation using Stbd1 (a glycogen-binding protein25) immunofluorescence revealed minimal glycogen foci in glucose-deprived BMDM, with no rChi3l1-dependent differences. High-glucose conditions, however, triggered robust glycogen aggregation, which was significantly attenuated by rChi3l1 (Figure 6B). Concordantly, extracellular acidification rate (ECAR) measurements showed reduced basal and total glycolytic capacity in rChi3l1-treated BMDMs (Figure 6C, D), confirming Chi3l1’s role in limiting glucose metabolism.

Earlier version: To test whether Chi3l1-glucose binding influence cell survival, we employed a palmitic acid (PA)-induced lipotoxicity cell-based model to better mimic the in vivo environment. rChi3l1 supplementation reduced PA-induced cleavage of caspase-3 (Figure 6F) and decreased KCs death (calcein/PI staining, Figure 6G). These findings collectively establish that Chi3l1 safeguards KCs by modulating glucose uptake, thereby suppressing metabolic stress-induced death.

Revised version: To test whether Chi3l1-glucose binding influence cell survival, we employed a palmitic acid (PA)-induced lipotoxicity cell-based model to better mimic the in vivo environment. rChi3l1 supplementation reduced PA-induced cleavage of caspase-3 (Figure 6E) and decreased KCs death (calcein/PI staining, Figure 6F). To validate this mechanism in vivo, we intraperitoneally injected 2-NBDG into WT and Chil1-/- mice, with or without supplementation of rChi3l1, to assess glucose uptake by KCs. Chil1-/- KCs displayed markedly increased 2-NBDG uptake compared with WT controls, whereas rChi3l1 supplementation significantly reduced glucose uptake. These results demonstrate that serum Chi3l1 limits glucose uptake by KCs in vivo (Figure 6G, H).Collectively, these findings demonstrate that Chi3l1 protects KCs from metabolic stress–induced death by regulating glucose uptake.

Discussion

Our findings establish Chi3l1 as a critical metabolic regulator that controls hepatic macrophage fate through a novel glucose-dependent mechanism in MASLD. Using cell-specific knockout models, we uncovered a fundamental dichotomy in Chi3l1 function: selective ablation in KCs dramatically accelerated MASLD progression and metabolic dysfunction, whereas deletion in MoMFs produced minimal metabolic effects. Single-cell transcriptomics revealed the molecular basis for this cell-type specificity - KCs exhibit a glucose-hungry metabolic phenotype that renders them uniquely dependent on Chi3l1-mediated regulation, while MoMFs maintain a relatively glucose-independent metabolic program. At the mechanistic level, we demonstrate that Chi3l1 functions as a physiological glucose sensor, directly binding extracellular glucose to limit its cellular uptake. This interaction establishes a crucial metabolic safeguard that specifically protects glucose-dependent KCs from lethal metabolic stress while sparing glucose-independent MoMFs. Through this precise modulation of glucose availability, Chi3l1 maintains metabolic homeostasis and preserves KCs populations during chronic dietary challenge (Figure 7).

Differential regulation of KCs and MoMFs fate by Chi3l1-glucose interaction.

KCs maintain a high-glucose activation state, while MoMFs exhibit a relatively low-glucose metabolic program. Chi3l1-glucose binding inhibits glucose uptake in KCs, thereby delaying KCs death and alleviating MASLD progression and metabolic dysfunction. In contrast, although Chi3l1-glucose binding similarly inhibits glucose uptake in MoMFs, their low basal glucose metabolism renders them resistant to this metabolic perturbation, resulting in minimal impact on MASLD pathogenesis.

Added in the revised version: Analysis of publicly available scRNA-seq datasets, including those from the Liver Atlas and prior studies7-9, indicates that Chil1 transcripts are mainly detected in neutrophils. In contrast, our immunofluorescence data show that Chi3l1 protein is predominantly localized in Kupffer cells under normal conditions and in both KCs and MoMFs during MASLD progression. This discrepancy likely reflects differences in transcript versus protein abundance and detection sensitivity. While scRNA-seq captures relative mRNA levels per cell, tissue-based staining reflects both expression and cell prevalence, highlighting macrophages as a major contributor to total hepatic Chi3l1 protein. Moreover, environmental factors such as diet, microbiota, or disease stage may influence Chil1 expression patterns across immune cell types.

Earlier version: Our study reveals fundamental differences in metabolic requirements between hepatic macrophage subsets that provide new insights into MASLD pathogenesis. We demonstrate that KCs and MoMFs play stage-specific roles in disease progression, with KCs serving as critical regulators of early metabolic homeostasis while MoMFs appear more involved in later inflammatory phases. This temporal specialization explains the striking dichotomy observed in our genetic models - KCs-specific Chi3l1 deletion dramatically exacerbated metabolic dysfunction, whereas MoMFs deletion showed minimal effects. The heightened glucose metabolism of KCs during MASLD renders them uniquely vulnerable to dietary stress. Chi3l1 serves as a crucial metabolic buffer in this context, directly protecting KCs through glucose modulation as evidenced by reduced glycogen accumulation and attenuated glycolytic flux. Our findings using the HFHC model complement previous findings in fibrogenic CDAA-HFAT models24 or MCD/CCL4 models25 or human livers26, collectively suggesting Chi3l1 may have dual roles in MASLD - maintaining metabolic balance through KCs in early disease while potentially influencing fibrogenesis via MoMFs in advanced stages. The accelerated KCs death in knockout models provides direct experimental evidence linking macrophage survival to metabolic outcomes, resolving key questions about MASLD progression mechanisms.

Revised version: Our study reveals fundamental differences in metabolic requirements between hepatic macrophage subsets that provide new insights into MASLD pathogenesis. We demonstrate that KCs and MoMFs play stage-specific roles in disease progression, with KCs serving as critical regulators of early metabolic homeostasis while MoMFs appear more involved in later inflammatory phases. This temporal specialization explains the striking dichotomy observed in our genetic models - KCs-specific Chi3l1 deletion dramatically exacerbated metabolic dysfunction, whereas MoMFs deletion showed minimal effects. The heightened glucose metabolism of KCs during MASLD renders them uniquely vulnerable to dietary stress. Chi3l1 serves as a crucial metabolic buffer in this context, directly protecting KCs through glucose modulation as evidenced by reduced glycogen accumulation and attenuated glycolytic flux. Our findings using the HFHC model complement previous findings in fibrogenic CDAA-HFAT models26 or MCD/CCL4 models27 or human livers28, collectively suggesting Chi3l1 may have dual roles in MASLD - maintaining metabolic balance through KCs in early disease while potentially influencing fibrogenesis via MoMFs in advanced stages. The accelerated KCs death in knockout models provides direct experimental evidence linking macrophage survival to metabolic outcomes, resolving key questions about MASLD progression mechanisms.

Added in the revised version: The structural characteristics of Chi3l1 have been extensively studied. Chi3l1 forms a homodimer, with each subunit containing a catalytic domain and a carbohydrate-binding domain. While the catalytic domain retains structural similarity to chitinases, it lacks enzymatic activity,29,30 and the carbohydrate-binding domain mediates interactions with carbohydrate ligands.29 While chitin-binding domains are traditionally known to interact with complex polysaccharides, our findings reveal that Chi3l1 (YKL-40), a mammalian chitinase-like protein, specifically binds to glucose—a simple monosaccharide. This represents a fundamental departure from canonical binding to insoluble polymers such as chitin and suggests a previously unrecognized role for Chi3l1 in monosaccharide recognition, potentially linking it to glucose metabolism and energy sensing. Furthermore, we observed that Chi3l1 protein levels increased in the serum of mice fed a high-fat, high-cholesterol (HFHC) diet for 16 weeks (Figure 5F) but plateaued with prolonged feeding (24 weeks; data not shown), suggesting an adaptive regulatory limit. Together, these findings indicate that Chi3l1 possesses glucose-binding capacity that may be functionally relevant but limited in vivo.

Removed in the revised version: The structural characteristics of Chi3l1 have been extensively studied. Chi3l1 is a homodimeric protein composed of two identical subunits. Each subunit consists of a catalytic domain and a carbohydrate-binding domain. The catalytic domain contains an active site capable of hydrolyzing chitin, a polymer of N-acetylglucosamine27,28. The carbohydrate-binding domain is responsible for the interaction between Chi3l1 and various carbohydrate ligands27. It is likely that Chi3l1 also interacts with glucose through its carbohydrate-binding domain. However, the precise molecular interactions between Chi3l1 and glucose are still not fully elucidated, and further research is needed to explore these interactions in greater detail. Furthermore, our data demonstrated that Chi3l1 is upregulated in the serum of mice fed a high-fat, high-cholesterol (HFHC) diet for 16 weeks (Fig. 3H). However, continuing to feed the mice for 24 weeks did not further increase the level of Chi3l1 (data not shown). This suggests that the level of Chi3l1 initially increases but reaches a plateau stage over time. These findings indicate that while Chi3l1 has the capability to bind with glucose, this interaction may be limited in living organisms.

Earlier version: Our findings carry important translational potential for MASLD treatment. The discovery of Chi3l1’s glucose-sensing function in KCs suggests two complementary therapeutic strategies: first, developing Chi3l1-based interventions to preserve KC populations during early metabolic dysfunction; second, creating cell-type-specific approaches that selectively modulate glucose metabolism in KCs while sparing MoMFs. The structural mapping of Chi3l1’s glucose-binding domain now enables rational design of small molecule mimetics to therapeutically enhance this protective pathway. Besides, several key questions emerge for future research to advance these therapeutic possibilities: (1) How glucose levels are coordinated with other death inducers such as lipid toxicity; (2) Whether competing carbohydrate ligands modulate Chi3l1’s glucose-sensing capacity in different metabolic states; (3) The clinical relevance of human Chi3l1 variants in MASLD susceptibility and progression. Addressing these questions will be crucial for translating our mechanistic insights into targeted therapies that account for the complex metabolic specialization of hepatic macrophage subsets.

Revised version: Our findings carry important translational potential for MASLD treatment. The discovery of Chi3l1’s glucose-sensing function in KCs suggests two complementary therapeutic strategies: first, developing Chi3l1-based interventions to preserve KC populations during early metabolic dysfunction; second, creating cell-type-specific approaches that selectively modulate glucose metabolism in KCs while sparing MoMFs. Importantly, although access to early-stage human liver tissue is limited due to the asymptomatic nature of the disease, multiple human studies have consistently reported elevated Chi3l1 levels in steatotic and fibrotic liver disease27,31,32, underscoring the clinical relevance of our mechanistic findings. Building on this evidence, the structural mapping of Chi3l1’s glucose-binding domain now enables rational design of small-molecule mimetics or biologics to therapeutically enhance this protective pathway. Besides, several key questions emerge for future research to advance these therapeutic possibilities: (1) How glucose levels are coordinated with other death inducers such as lipid toxicity; (2) Whether competing carbohydrate ligands modulate Chi3l1’s glucose-sensing capacity in different metabolic states; (3) The clinical relevance of human Chi3l1 variants in MASLD susceptibility and progression. Addressing these questions will be crucial for translating our mechanistic insights into targeted therapies that account for the complex metabolic specialization of hepatic macrophage subsets.

Our findings reveal a novel metabolic checkpoint in which Chi3l1 selectively sustains KCs populations by modulating glucose metabolism, offering key insights into MASLD pathogenesis. The study highlights the therapeutic potential of targeting Chi3l1-glucose interactions to preserve protective KCs and curb MASLD progression. Future research should explore whether Chi3l1 supplementation or pharmacological modulation can rescue KCs viability, as well as investigate whether this mechanism extends to other macrophage-driven metabolic disorders, such as MASH or diabetes. By identifying cell type-specific metabolic vulnerabilities, this work paves the way for precision therapies that selectively manipulate macrophage subsets to treat liver disease.